Before you turn this notebook in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = ""
```

_Understanding Deep Learning_

---

<a href="https://colab.research.google.com/github/DL4DS/fa2024/blob/main/static_files/assignments/01_notebook.ipynb" target="_blank"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Notebook 1.1 -- Background Mathematics

The purpose of this Python notebook is to make sure you can run notebooks on
either your local environment or on CoLab and to familiarize yourself with some of the background mathematical concepts that you are going to need to understand deep learning.

It's not meant to be difficult and it may be that you know some or all of this information already.

Math is *NOT* a spectator sport.  You won't learn it by just listening to lectures or reading books.  It really helps to interact with it and explore yourself.

Work through the cells below, running each cell in turn.  In various places you will see the words **YOUR CODE HERE**. Follow the instructions at these places and write code to complete the functions.  There are also questions interspersed in the text.

```python
# Imports math library
import numpy as np
# Imports plotting library
import matplotlib.pyplot as plt
```

### Linear functions
We will be using the term *linear equation* to mean a weighted sum of inputs plus an offset. If there is just one input $x$, then this is a straight line:

$$ y=\beta+\omega x,$$

where $\beta$ is the y-intercept of the linear and $\omega$ is the slope of the line. When there are two inputs $x_{1}$ and $x_{2}$, then this becomes:

$$y=\beta+\omega_1 x_1 + \omega_2 x_2.$$

Any other functions are by definition **non-linear**.




<!-- grade_id: linear_func_1d, points: 0, grade: false, solution: true -->
```python
# Define a linear function with just one input, x
def linear_function_1D(x,beta,omega):
  # replace the code line below with formula for 1D linear equation

  # YOUR CODE HERE
  raise NotImplementedError()

  return y
```

<!-- grade_id: linear_func_1d_tests, points: 2, grade: true, solution: false -->
```python
assert linear_function_1D(2.5, -1.0, 0.3) == -0.25
```

```python
# Plot the 1D linear function

# Define an array of x values from 0 to 10 with increments of 0.01
# https://numpy.org/doc/stable/reference/generated/numpy.arange.html
x = np.arange(0.0,10.0, 0.01)
# Compute y using the function you filled in above
beta = 0.0; omega = 1.0

y = linear_function_1D(x,beta,omega)

# Plot this function
fig, ax = plt.subplots()
ax.plot(x,y,'r-')
ax.set_ylim([0,10]);ax.set_xlim([0,10])
ax.set_xlabel('x'); ax.set_ylabel('y')
plt.show
```

### Optional: 

Experiment with changing the values of beta and omega to understand what they do.  

<!-- grade_id: cell-27f75ceb132dc6ed, points: 0, grade: false, solution: true -->
```python
# Set beta and omega above so that the line crosses the y-axis at y=10 and the x-axis at x=5

# Copy your beta and omega assignments below as well

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: beta_omega, points: 1, grade: true, solution: false -->
```python
# Don't edit this cell.

```

Now let's investigate a 2D linear function

```python
# Code to draw 2D function -- read it so you know what is going on, but you don't have to change it
def draw_2D_function(x1_mesh, x2_mesh, y):
    fig, ax = plt.subplots()
    fig.set_size_inches(7,7)
    pos = ax.contourf(x1_mesh, x2_mesh, y, levels=256 ,cmap = 'hot', vmin=-10,vmax=10.0)
    fig.colorbar(pos, ax=ax)
    ax.set_xlabel('x1');ax.set_ylabel('x2')
    levels = np.arange(-10,10,1.0)
    ax.contour(x1_mesh, x2_mesh, y, levels, cmap='winter')
    plt.show()
```

<!-- grade_id: linear_func_2d, points: 0, grade: false, solution: true -->
```python
# Define a linear function with two inputs, x1 and x2
def linear_function_2D(x1,x2,beta,omega1,omega2):
  # TODO -- write the formula for the 2D linear equation

  # YOUR CODE HERE
  raise NotImplementedError()

  return y
```

<!-- grade_id: linear_func_2d_tests, points: 2, grade: true, solution: false -->
```python
assert linear_function_2D(1.0,2.0,0.0,1.0,1.0) == 3.0
```

```python
# Plot the 2D function

# Make 2D array of x and y points
x1 = np.arange(0.0, 10.0, 0.1)
x2 = np.arange(0.0, 10.0, 0.1)
x1,x2 = np.meshgrid(x1,x2)  # https://www.geeksforgeeks.org/numpy-meshgrid-function/

# Compute the 2D function for given values of omega1, omega2
beta = 0.0; omega1 = 1.0; omega2 = -0.5
y  = linear_function_2D(x1,x2,beta, omega1, omega2)

# Draw the function.
# Color represents y value (brighter = higher value)
# Black = -10 or less, White = +10 or more
# 0 = mid orange
# Lines are contours where value is equal
draw_2D_function(x1,x2,y)
```

### Question: Predict what this plot will look like if you set omega_1 to zero

Change the code and see if you are right.

<!-- grade_id: cell-9b2b4ef013b0fb4a, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Question: Predict what this plot will look like if you set omega_2 to zero

Change the code and see if you are right.

<!-- grade_id: cell-d6d96ebbaed5be01, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Question: Predict what this plot will look like if you set beta to -5

Change the code and see if you are correct

<!-- grade_id: cell-b5a74c0ca60a405d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Often we will want to compute many linear functions at the same time.  For example, we might have
three inputs, $x_1$, $x_2$, and $x_3$ and want to compute two linear functions giving $y_1$ and $y_2$. 
Of course, we could do this by just running each equation separately,

$$
\begin{align}
y_1 &=& \beta_1 + \omega_{11} x_1 + \omega_{12} x_2 + \omega_{13} x_3\\
y_2 &=& \beta_2 + \omega_{21} x_1 + \omega_{22} x_2 + \omega_{23} x_3.
\end{align}
$$

However, we can write it more compactly with vectors and matrices:

$$
\begin{bmatrix} y_1\\ y_2 \end{bmatrix} = \begin{bmatrix}\beta_{1}\\\beta_{2}\end{bmatrix}+ \begin{bmatrix}\omega_{11}&\omega_{12}&\omega_{13}\\\omega_{21}&\omega_{22}&\omega_{23}\end{bmatrix}\begin{bmatrix}x_{1}\\x_{2}\\x_{3}\end{bmatrix},
$$

or

$$
\mathbf{y} = \boldsymbol\beta +\boldsymbol\Omega\mathbf{x}.
$$

for short.  Here, lowercase bold symbols are used for vectors.  Upper case bold symbols are used for matrices.



<!-- grade_id: cell-b4cd1ed88f18ba63, points: 0, grade: false, solution: true -->
```python
# Define a linear function with three inputs, x1, x2, and x_3
def linear_function_3D(x1,x2,x3,beta,omega1,omega2,omega3):
  # TODO -- write the formula for a single 3D linear equation

  # YOUR CODE HERE
  raise NotImplementedError()

  return y
```

<!-- grade_id: cell-44cf5af9dd4ec56c, points: 2, grade: true, solution: false -->
```python
assert linear_function_3D(1.0,2.0,3.0,0.0,1.0,1.0,1.0) == 6.0
```

Let's compute two linear equations, using both the individual equations and the vector / matrix form and check they give the same result

```python
# Define the parameters
beta1 = 0.5; beta2 = 0.2
omega11 =  -1.0 ; omega12 = 0.4; omega13 = -0.3
omega21 =  0.1  ; omega22 = 0.1; omega23 = 1.2

# Define the inputs
x1 = 4 ; x2 =-1; x3 = 2

# Compute using the individual equations
y1 = linear_function_3D(x1,x2,x3,beta1,omega11,omega12,omega13)
y2 = linear_function_3D(x1,x2,x3,beta2,omega21,omega22,omega23)
print("Individual equations")
print('y1 = %3.3f\ny2 = %3.3f'%((y1,y2)))

# Define vectors and matrices
beta_vec = np.array([[beta1],[beta2]])
omega_mat = np.array([[omega11,omega12,omega13],[omega21,omega22,omega23]])
x_vec = np.array([[x1], [x2], [x3]])

# Compute with vector/matrix form
y_vec = beta_vec+np.matmul(omega_mat, x_vec)
print("Matrix/vector form")
print('y1= %3.3f\ny2 = %3.3f'%((y_vec[0][0],y_vec[1][0])))

```

# Questions

1.  A single linear equation with three inputs (i.e. **linear_function_3D()**) associates a value y with each point in a 3D space ($x_1$,$x_2$,$x_3$).  Is it possible to visualize this?   What value is at position (0,0,0)?

<!-- grade_id: cell-0e2b005301d2f37e, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

2.  Write code to compute three linear equations with two inputs ($x_1$, $x_2$) using both the individual equations and the matrix form (you can make up any values for the inputs $\beta_{i}$ and the slopes $\omega_{ij}$.

<!-- grade_id: cell-4c4e2dde94244d9e, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

# Special functions

Throughout the book and class, we'll be using some special functions (see Appendix B.1.3).  The most important of these are the logarithm and exponential functions.  Let's investigate their properties.

We'll start with the exponential function $y=\exp[x]=e^x$ which maps the real
line $[-\infty,+\infty]$ to non-negative numbers $[0,+\infty]$.

```python
# Draw the exponential function

# Define an array of x values from -5 to 5 with increments of 0.01
x = np.arange(-5.0,5.0, 0.01)
y = np.exp(x) ;

# Plot this function
fig, ax = plt.subplots()
ax.plot(x,y,'r-')
ax.set_ylim([0,100]);ax.set_xlim([-5,5])
ax.set_xlabel('x'); ax.set_ylabel('exp[x]')
plt.show
```

# Questions

1. What is $\exp[0]$?

<!-- grade_id: cell-a2f5bc5c9b439e6e, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

2. What is $\exp[1]$?

<!-- grade_id: cell-ce8a1c245304a02c, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

3. What is $\exp[-\infty]$?

<!-- grade_id: cell-8128f779d2fecc83, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

4. What is $\exp[+\infty]$?

<!-- grade_id: cell-98b05cf3da43c8e2, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

5. A function is convex if we can draw a straight line between any two points on the function, and this line always lies above the function. Similarly, a function is concave if a straight line between any two points always lies below the function.  Is the exponential function convex or concave or neither?

<!-- grade_id: cell-0a5cd86079f44ef5, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Now let's consider the logarithm function $y=\log[x]$. Throughout the book we always use natural (base $e$) logarithms. The log function maps non-negative numbers $[0,\infty]$ to real numbers $[-\infty,\infty]$.  It is the inverse of the exponential function.  So when we compute $\log[x]$ we are really asking "What is the number $y$ so that $e^y=x$?"

```python
# Draw the logarithm function

# Define an array of x values from -5 to 5 with increments of 0.01
x = np.arange(0.01,5.0, 0.01)
y = np.log(x) ;

# Plot this function
fig, ax = plt.subplots()
ax.plot(x,y,'r-')
ax.set_ylim([-5,5]);ax.set_xlim([0,5])
ax.set_xlabel('x'); ax.set_ylabel('$\log[x]$')
plt.show
```

# Questions

1. What is $\log[0]$?  

<!-- grade_id: cell-6d16991c50c8f707, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

2. What is $\log[1]$?

<!-- grade_id: cell-168daf8582c010e4, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

3. What is $\log[e]$?

<!-- grade_id: cell-bba5c4164f473d21, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

4. What is $\log[\exp[3]]$?

<!-- grade_id: cell-d8f7e4377db60f29, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

5. What is $\exp[\log[4]]$?

<!-- grade_id: cell-20eb8b5e640e99d8, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

6. What is $\log[-1]$?

<!-- grade_id: cell-b1c4d7e1373f3a17, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

7. Is the logarithm function concave or convex?

<!-- grade_id: cell-dd9689a00c350ff7, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE